## Remember that the solutions need to be saved int variable provided.

```python
# Import the needed libraries
import numpy as np
```

**13. Create a 10x10 array with random values and find the minimum and maximum values**

<!-- grade_id: cell-d37d68df72162149, points: 0, grade: false, solution: true -->
```python
# 1 point
sol1 = np.random.rand(10,10)
min = sol1.min()
max = sol1.max()
```

<!-- grade_id: cell-f1d4fa08074f4dcd, points: 1, grade: true, solution: false -->
```python
assert sol1.shape == (10,10)
```

**15. Create a 2d 10x10 array with 1 on the border and 0 inside**

<!-- grade_id: cell-6584e6c544771f0c, points: 0, grade: false, solution: true -->
```python
#1 points
sol2 = np.ones((10,10),dtype='int32')
sol2[1:-1,1:-1] = 0
sol2
```

Output:
```
array([[1, 1, 1, 1, 1, 1, 1, 1, 1, 1],
       [1, 0, 0, 0, 0, 0, 0, 0, 0, 1],
       [1, 0, 0, 0, 0, 0, 0, 0, 0, 1],
       [1, 0, 0, 0, 0, 0, 0, 0, 0, 1],
       [1, 0, 0, 0, 0, 0, 0, 0, 0, 1],
       [1, 0, 0, 0, 0, 0, 0, 0, 0, 1],
       [1, 0, 0, 0, 0, 0, 0, 0, 0, 1],
       [1, 0, 0, 0, 0, 0, 0, 0, 0, 1],
       [1, 0, 0, 0, 0, 0, 0, 0, 0, 1],
       [1, 1, 1, 1, 1, 1, 1, 1, 1, 1]])
```

<!-- grade_id: cell-9429aab0f4477a86, points: 1, grade: true, solution: false -->
```python
assert sol2.shape == (10,10)
assert sol2[0].all() == 1
assert sol2[:,0].all() == 1
assert sol2[1][1] == 0
assert sol2[2][2] == 0
assert sol2[3][3] == 0
assert sol2[4][4] == 0
assert sol2[1][2] == 0
assert sol2[2][3] == 0
assert sol2[3][4] == 0
assert sol2[4][5] == 0
```

**34. How to get all the dates corresponding to the month of July 2016**

<!-- grade_id: cell-9784b014be5fdf5a, points: 0, grade: false, solution: true -->
```python
#2 points
sol3 = np.arange("2016-07", "2016-08", dtype='datetime64[D]')

```

<!-- grade_id: cell-8beeae591e9b50ed, points: 2, grade: true, solution: false -->
```python
assert sol3.shape == (31,)
assert type(sol3[0]) is np.datetime64
```

**42. Consider two random arrays based on the normal and uniform distributions 10x10 sol41 and sol42, check if they are equal**

<!-- grade_id: cell-5000744434ae05a2, points: 0, grade: false, solution: true -->
```python
#2 points
sol41 = np.random.random((10,10))
sol42 = np.random.uniform((10,10))
comparisson = np.ones((10,10),dtype='bool')
if sol41 != sol42:
    comparisson[:] = 0
comparisson
```

Output:
```
[[0.37694497 0.87574229 0.88682458 0.79404443 0.24987289 0.08592235
  0.39102928 0.796829   0.97300644 0.43889434]
 [0.22422321 0.14331885 0.34006699 0.96719792 0.54084044 0.59544312
  0.73312385 0.27626848 0.01744296 0.1792946 ]
 [0.85070205 0.24220502 0.43283778 0.58542271 0.78513998 0.87729828
  0.94439597 0.59852332 0.03523319 0.0575719 ]
 [0.49261745 0.68013077 0.56295514 0.09588235 0.43687302 0.8008063
  0.97996877 0.53619089 0.07681174 0.5686293 ]
 [0.56896358 0.22447353 0.10146245 0.11860519 0.02669508 0.81824788
  0.07553013 0.3874986  0.58226845 0.8056386 ]
 [0.72649386 0.98612178 0.03339659 0.74609589 0.92603215 0.33318095
  0.77253245 0.79880286 0.37137572 0.26513019]
 [0.24967085 0.89200751 0.7450441  0.60982344 0.11157358 0.3125394
  0.77124274 0.22048092 0.68239649 0.07222066]
 [0.10802376 0.52152296 0.81549538 0.16263444 0.61047467 0.23885884
  0.56998636 0.68346025 0.36128325 0.00803387]
 [0.71101448 0.75318278 0.7778592  0.82808078 0.00682271 0.91942857
  0.68152332 0
```

<!-- grade_id: cell-807bf6c3ff3f3b0d, points: 2, grade: true, solution: false -->
```python
bol = True

for l in comparisson:
    for c in l:
        if c:
            bol = False


assert bol
```

**72. How to swap two rows of an array?** 

<!-- grade_id: cell-079026f268b5211c, points: 0, grade: false, solution: true -->
```python
#3 points
sol5 = np.array([[1,2,3],[4,5,6]])
tmp = np.copy(sol5)
tmp[0] = sol5[1]
tmp[1] = sol5[0]
sol5 = tmp
```

<!-- grade_id: cell-e4f878f0864f33bd, points: 3, grade: true, solution: false -->
```python
assert sol5[0][0] == 4
assert sol5[0][1] == 5
assert sol5[0][2] == 6

assert sol5[1][0] == 1
assert sol5[1][1] == 2
assert sol5[1][2] == 3

```

**75. How to compute averages using a sliding window over an array?**

<!-- grade_id: cell-ccafbda4e4cb7996, points: 0, grade: false, solution: true -->
```python
arr = np.array([1,2,3,4,5,6,7,8,9,10])
sol6 = np.cumsum(arr) / np.size(arr)
sol6

```

Output:
```
array([0.1, 0.3, 0.6, 1. , 1.5, 2.1, 2.8, 3.6, 4.5, 5.5])
```

<!-- grade_id: cell-1983fc0078a860ff, points: 3, grade: true, solution: false -->
```python
assert sol6[-1] == 5.5
assert sol6[0] == 0.1
assert sol6[4] == 1.5
```

```python

```